---

# Numerical Analysis
# Differentiation

## Remember

### Forward Pass Difference Formula

Given a function $f(x)$, we want to calculate its derivative $f'(x_0)$ at point $x_0$. For a small step $h$, the derivative can be approximated using the following forward difference formula

$$f'(x) = \frac{f(x_0 + h) - f(x_0)}{h}$$

### Three Point Mid-point Formula

Given the same conditions as the forward pass difference formula, the three point mid-point formula is as follows:

$$f'(x) = \frac{f(x_0 + h) - f(x_0 - h)}{2h}$$

#### TO-DO (1)

Deﬁne a function `df_approx` that takes in the function `f`, an initial point `x0` and the step size `h` and returns the approximation of the function $f(x)$ at point $x_0$ using the forward pass difference formula

<!-- grade_id: cell-db97c555e129ba01, points: 0, grade: false, solution: true -->
```python
import numpy as np

def df_approx(f, x0, h):
    # YOUR CODE HERE
    return (f(x0+h)-f(x0))/(h)
    raise NotImplementedError()
```

<!-- grade_id: cell-05aa7e9e5bf73475, points: 2, grade: true, solution: false -->
```python
import math
assert np.isclose(df_approx(math.sin, 0, 1e-8), 1.0)
assert np.isclose(df_approx(math.cos, 0, 1e-8), 0.0)
```

#### TO-DO (2)

Deﬁne a function `df_midpoint_approx` that takes in the function `f`, an initial point `x0` and the step size `h` and returns the approximation of the function $f(x)$ at point $x_0$ using the three point mid-point formula.

<!-- grade_id: cell-1a2a21e628de402b, points: 0, grade: false, solution: true -->
```python
import numpy as np

def df_midpoint_approx(f, x0, h):
    # YOUR CODE HERE
    return (f(x0+h)-f(x0-h))/(2*h)
    raise NotImplementedError()
```

<!-- grade_id: cell-590a3e49a220661b, points: 2, grade: true, solution: false -->
```python
import math
assert np.isclose(df_midpoint_approx(math.sin, 0, 1e-8), 1.0)
assert np.isclose(df_midpoint_approx(math.cos, 0, 1e-8), 0.0)
```

For $x=0.7$, and step size $h = 2^i$ with $i$ varying from $0$ to $-30$ with step size 1: 

* Calculate the approximation of the derivative of the function $f(x) = cos(x)$ using the forward pass difference formula.
* Calculate the error and the error ratio for each choice of $h$ where the error is computed as `error = abs(diff. approx - true approx)` and the error ratio is computed as `error[i-1] / error[i]`.
* After finishing, print a table with the following three columns: i (where h=2^(-i)) , approximation of f'(x) , Error Ration.

You should calculate the approximate derivatives in a list names `approx_deriv` where each element in the list corresponds to the approximate derivative at one value of $h$. **Additionally**, store the error ratio in a list `error_ratios` where each element corresponds to the error ratio compute at a value of $h$.

Verify that the error ratios are approximately 2, since $E_h / E_{h/2} \approx E/(E/2) = 2$


<!-- grade_id: cell-dec0d54d28bda401, points: 0, grade: false, solution: true -->
```python
import numpy as np 
x = 0.7
MAX_POWER = 30
h = [2**(-i) for i in range(MAX_POWER)]
f= np.cos
T_value = -np.sin(x)
approx_deriv= []
error_ratios= []
d0 = df_approx(f,x,h[0])
approx_deriv.append(d0)
for hi in h[1::]:
  d = df_approx(f,x,hi)
  approx_deriv.append(d)
  error0 = abs(d0 - T_value)
  error  = abs(d - T_value)
  error_r = error0/error
  error_ratios.append(error_r)
  d0 = d


# YOUR CODE HERE
#raise NotImplementedError()
```

<!-- grade_id: cell-eefa4aae2600f867, points: 4, grade: true, solution: false -->
```python
assert np.allclose(np.array(approx_deriv), np.array([-0.8936866815800131, -0.8049688656156297, -0.7326363912824196, -0.6902817732851885, -0.6676918955783648, -0.6560625258790012, -0.6501666822936372, -0.647198783534563, -0.6457098794061835, -0.6449641936163175, -0.6445910429116566, -0.6444043906847128, -0.6443110453624286, -0.6442643679001776, -0.6442410279687465, -0.6442293577056262, -0.644223522496759, -0.6442206048814114, -0.6442191460810136, -0.6442184166517109, -0.6442180520389229, -0.6442178697325289, -0.6442177784629166, -0.6442177332937717, -0.64421771094203, -0.6442176997661591, -0.6442176923155785, -0.6442176997661591, -0.6442176997661591, -0.6442177295684814]))
assert np.allclose(np.array(error_ratios), np.array([1.5518952760383549, 1.8180675696922572, 1.919471580388214, 1.9623275630421197, 1.9818090437133529, 1.9910654707085305, 1.9955729246949976, 1.9977965035720624, 1.9989007612205443, 1.9994510079290797, 1.9997256607803775, 1.9998628672827814, 1.9999314440828047, 1.9999657412458582, 1.9999824501611332, 1.999991397000223, 1.9999902754200525, 1.9999705761583981, 2.0000209522790238, 1.9994834339616852, 1.9989674013287748, 2.0004865618718357, 1.9807422636904461, 1.9429388371031604, 1.892038105457464, 2.467259894846999, 0.40530792969502405, 1.0, 0.2959658422148293]))
```

Repeat using the three point mid-point formula.

<!-- grade_id: cell-19d4a4d37d276fdb, points: 0, grade: false, solution: true -->
```python
import numpy as np 
x = 0.7
MAX_POWER = 30
h = [2**(-i) for i in range(MAX_POWER)]
f= np.cos
T_value = -np.sin(x)
approx_deriv= []
error_ratios= []
d0 = df_midpoint_approx(f,x,h[0])
approx_deriv.append(d0)
for hi in h[1::]:
  d = df_midpoint_approx(f,x,hi)
  approx_deriv.append(d)
  error0 = abs(d0 - T_value)
  error  = abs(d - T_value)
  error_r = error0/error
  error_ratios.append(error_r)
  d0 = d
print(error_ratios)
# YOUR CODE HERE
#raise NotImplementedError()
```

Output:
```
[3.852567805845165, 3.962661493592123, 3.990635109130402, 3.9976568820714045, 3.999414102034443, 3.9998535181303843, 3.9999633814349256, 3.999990854900724, 3.999997768616538, 3.9999999566300595, 3.9999987162501736, 4.000012629362491, 3.9996242076860353, 4.000202061022429, 4.042175649598349, 3.7774199016782504, 3.271769273433167, 10.953827460510327, 0.04782937176730401, 1.1057772636720005, 0.19188132729940688, 1.622751375687844, 1.0, 0.10536291436245686, 0.8259508274083478, 0.21496585143079605, 0.4672598948469992, 2.140136594276816, 0.1893841406099519]
```

<!-- grade_id: cell-9143196205ed5154, points: 2, grade: true, solution: false -->
```python
assert np.allclose(np.array(approx_deriv), np.array([-0.5420904917105653, -0.617708823364568, -0.6375280257775868, -0.6425413471872568, -0.6437983565897927, -0.6441128392181561, -0.6441914742728763, -0.6442111339364942, -0.6442160489086461, -0.6442172776552013, -0.6442175848420106, -0.6442176616387769, -0.6442176808379827, -0.6442176856376136, -0.6442176868367824, -0.6442176871387346, -0.6442176872114942, -0.64421768723696, -0.64421768723696, -0.6442176872224081, -0.644217687251512, -0.6442176873097196, -0.6442176871933043, -0.6442176871933043, -0.6442176876589656, -0.644217686727643, -0.6442176848649979, -0.6442176923155785, -0.6442176848649979, -0.6442176997661591]))
assert np.allclose(np.array(error_ratios), np.array([3.852567805845165, 3.962661493592123, 3.990635109130402, 3.9976568820714045, 3.999414102034443, 3.9998535181303843, 3.9999633814349256, 3.999990854900724, 3.999997768616538, 3.9999999566300595, 3.9999987162501736, 4.000012629362491, 3.9996242076860353, 4.000202061022429, 4.042175649598349, 3.7774199016782504, 3.271769273433167, 10.953827460510327, 0.04782937176730401, 1.1057772636720005, 0.19188132729940688, 1.622751375687844, 1.0, 0.10536291436245686, 0.8259508274083478, 0.21496585143079605, 0.4672598948469992, 2.140136594276816, 0.1893841406099519]
))
```

**\[THE END\]**

```python

```