---

# Numerical Analysis
# Lagrange Interpolation

## Remember

If $x_0, x_1, \dots , x_n$ are $n + 1$ distinct numbers and $f$ is a function whose values are given at these numbers, then a unique polynomial $P(x)$ of degree at most $n$ exists with :

$ f(x_k) = P(x_k) $ for each $k = 0, 1, \dots k$

This polynomial is given by:

$$ P(x) = f(x_0)L_{0}(x)+\dots+f(x_n)L_{n}(x) = \sum_{k=0}^{n}f(x_k)L_{k}(x) $$

where for each $k=0,1,\dots, n$, $L_{k}(x)$ is defined as follows:

$$ L_{k}(x) = \frac{(x-x_0)(x-x_1)\dots (x-x_{k-1})(x-x_{k+1})\dots (x-x_n)}{(x_k-x_0)(x_k-x_1)\dots (x_k-x_{k-1})(x_k-x_{k+1})\dots (x_k-x_n)} $$

$$ = \prod_{i=0\\ i\neq k}^{n} \frac{(x-x_i)}{(x_k-x_i)}$$

It is required to estimate the value of the function $\sqrt{u}$ at the value $u=1.4$ by interpolating the function at the following list of values

| $x$ | $1.0$ | $2.0$ | $3.0$ |
|---|-------|-------|-------|
| $y$ | $1.0$ | $1.4142$| $1.7321$|

Using Lagrnage Interpolation in Python.

Lagrange Algorithm Pseudo-Code:

```
Input:
    x: list of x values
    y: list of y values (list of f(x) values)
    u: the x value at which we want to interpolate f(x)
Output:
    An estimation of the value of f(u)

Algorithm:
    Calculate L_k(x) as follows:
        numenator = empty list
        denominator = empty list
        Repeat for i = 0, 1, ... , n (the number of x values)
            add to numenator (u - x_i) , where i != k
            add to denominator (x_k - x_i) , where i != k
        return Product(numenator) / Product(denominator)
     
    Calculate f(u) as follows:
        return sum of (y_i * L_i) for i = 0, 1, ..., n        

```

#### TO-DO (1)

Deﬁne a function `product(x)` that takes in a list of values and return their product (multiplication).

<!-- grade_id: cell-db97c555e129ba01, points: 0, grade: false, solution: true -->
```python
def product(x):
    # YOUR CODE HERE
    p =1
    for _ in x:
      p *= _
    return p
    raise NotImplementedError()
```

<!-- grade_id: cell-731c6a3f26924bdb, points: 2, grade: true, solution: false -->
```python
assert product([1,2,3,4]) == 24
assert product([4,3,2,1]) == 24
```

#### TO-DO (2)

Defien a function `L(k, x, u)` that takes in the index of the current term $L_k$, the x values `x`, and the input x value `u` - and returns the value of $L_k(x)$. Check the **Remember** cell above for how to compute this function.

<!-- grade_id: cell-1a2a21e628de402b, points: 0, grade: false, solution: true -->
```python
def L(k, x, u):
    # YOUR CODE HERE
    lst=[]
    for i in range(len(x)):
      if i!= k:
        lst.append((u-x[i])/(x[k]-x[i]))
    return product(lst)
    raise NotImplementedError()
```

<!-- grade_id: cell-590a3e49a220661b, points: 3, grade: true, solution: false -->
```python
import numpy as np
assert np.isclose(L(1, [1,2,3], 1.4), 0.639999999999)
```

#### TO-DO (3)

Define a function `lagrange` that takes in the x values `x`, the y values `y`, the input x value `u` and returns an estimation of the interpolated function at `u`, `f(u)`. Check the **Remember** cell above for how to compute this value.

<!-- grade_id: cell-14f505e9b940e351, points: 0, grade: false, solution: true -->
```python
def lagrange(x, y, u):
    # YOUR CODE HERE
    rslt= 0
    for i_ in range(len(x)):
      L_k = L(i_, x,u)
      rslt += L_k*y[i_]
    return rslt
    raise NotImplementedError()
```

<!-- grade_id: cell-e1e007d3f7c9885a, points: 1, grade: true, solution: false -->
```python
import numpy as np
assert np.isclose(lagrange([1,2,3], [1,1.4142,1.7321], 1.4), 1.177236)
```