---

# Numerical Analysis
# Solving a System of Linear Equations

## Remember

### Jacobi

To solve the system of equations $Ax = b$ using Jacobi's method, you decompose $A$ into three matrices $D+L+U$ where $D$ is the diagonal matrix of $A$, $L$ is the strict lower triangular matrix of $A$ and $U$ is the strict upper triangular matrix of $A$.
Then to find a solution, you iterate through the following formula
$$
x^{(k)} = D^{-1}(b - (L+U)x^{(k-1)})
$$

### Gauss-Siedel
To solve the system of equations $Ax = b$ using Gauss-Siedel method, you decompose $A$ into three matrices $D+L+U$ where $D$ is the diagonal matrix of $A$, $L$ is the strict lower triangular matrix of $A$ and $U$ is the strict upper triangular matrix of $A$.
Then to find a solution, you iterate through the following formula
$$
x^{(k)} = (D+L)^{-1}(b - Ux^{(k-1)})
$$

#### TO-DO (1)

Deﬁne a function `jacobi` that takes in $n\times n$ square matrix `A`, vector `b` of size $n$, and a number of iterations `N`. Start with an initial solution $x^{(0)} = 0$ and use the Jacobi's formula to find an approximate solution to the equation `Ax = b`.

<!-- grade_id: cell-db97c555e129ba01, points: 0, grade: false, solution: true -->
```python
import numpy as np
#from numpy import linalg
def jacobi(A, b, N):
    # YOUR CODE HERE
    X_0 = np.zeros_like(b)
    D= np.diag(np.diag(A))
    D_1= np.linalg.inv(D)
    for i in range(N):
      X_0 = D_1@(b-(A-D)@X_0)
    return X_0
    raise NotImplementedError()
```

<!-- grade_id: cell-05aa7e9e5bf73475, points: 5, grade: true, solution: false -->
```python
np.random.seed(42)
A = 10 * np.eye(10)
A = A + np.random.rand(10,10)
b = np.sum(A, axis=1)
s = jacobi(A,b,20)
print(s)
assert np.allclose(s, np.ones(10))
```

Output:
```
[0.99999999 0.99999999 0.99999999 0.99999999 0.99999999 0.99999999
 0.99999999 0.99999999 0.99999999 0.99999999]
```

#### TO-DO (2)

Deﬁne a function `siedel` that takes in $n\times n$ square matrix `A`, vector `b` of size $n$, and a number of iterations `N`. Start with an initial solution $x^{(0)} = 0$ and use the Gauss-Siedel formula to find an approximate solution to the equation `Ax = b`.

<!-- grade_id: cell-1a2a21e628de402b, points: 0, grade: false, solution: true -->
```python
import numpy as np
from numpy.linalg import inv

def siedel(A, b, N):
    # YOUR CODE HERE
    X_0 = np.zeros_like(b)
    D= np.diag(np.diag(A))
    L = np.tril(A,-1)
    U = np.triu(A,1)
    DL_1= np.linalg.inv(D+L)
    for i in range(N):
      X_0 = DL_1@(b-U@X_0)
    print(X_0)
    return X_0
    raise NotImplementedError()
```

<!-- grade_id: cell-590a3e49a220661b, points: 5, grade: true, solution: false -->
```python
np.random.seed(42)
A = 10 * np.eye(10)
A = A + np.random.rand(10,10)
b = np.sum(A, axis=1)
s = siedel(A,b,10)
assert np.allclose(s, np.ones(10))
```

Output:
```
[1. 1. 1. 1. 1. 1. 1. 1. 1. 1.]
```

**\[THE END\]**